# Assignment 0: Tools & setting up your computer

This assignment simply checks that you can access our assignments and will familiarize you with submitting assignments via NBGrader.


**PLEASE DO NOT CHANGE THE NAME OF THIS FILE.**

**PLEASE DO NOT COPY & PASTE OR DELETE CELLS INCLUDED IN THE ASSIGNMENT.**


## How to complete assignments

Whenever you see:

```
# YOUR CODE HERE
raise NotImplementedError()
```

You need to **replace (in other words, delete) these lines of code with code that answers the questions** and meets the specified criteria. Make sure you remove the 'raise' line when you do this (or your notebook will raise an error, regardless of any other code, and thus fail the grading tests).

You should write the answer to the questions in those cells (the ones with `# YOUR CODE HERE`), but you can also add extra cells to explore / investigate things if you need / want to. 

Any cell with `assert` statements in it is a test cell. You should not try to change or delete these cells. Note that there might be more than one assert that tests a particular question. 

If a test does fail, reading the error that is printed out should let you know which test failed, which may be useful for fixing it.

Note that some cells, including the test cells, may be read only, which means they won't let you edit them. If you cannot edit a cell - that is normal, and you shouldn't need to edit that cell.

## Q1

Create a variable called `my_name`, and assign a string to it that is your name. In other words, enter `my_name = 'name'`. Be sure to put your name in single quotes, e.g., `'Ashley'`.

<!-- grade_id: cell-4995bf24615505c5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-14c71d611db9757d, points: 5, grade: true, solution: false -->
```python
## TESTS FOR Q1, worth 5 points
# This is a test cell. It checks that the code you wrote does what it's supposed to.
# You can use the information in this cell to make sure you've met the expectations for the question.
assert  my_name                     # This tests that a variable "my_name" exists
assert  isinstance(my_name, str)    # isinstance() tests that "my_name" is of a particular type (here, a string)
```

## Q2
Occasionally, we'll use these notebooks to ask multiple choice questions. When you respond to a multiple choice question, **be sure to put your response in single quotation marks, the way you did in Q1**.

For example, if you wanted to respond "A", you should write `answer = 'A'`

Respond to the question below:

What is the **kernel**?

- A) a list of the variables stored in Python memory
- B) where your code runs; your Jupyter notebook is connected to the kernel
- C) an object that returns a string
- D) a way to specify that you are using math variables and not programming variables

<!-- grade_id: cell-8d20cdb4d0f455ff, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e986bfb4e19f48c3, points: 5, grade: true, solution: false -->
```python
## TESTS FOR Q2, worth 5 points [Note: Includes A Hidden Test]
assert answer

```

## Q3

Write code below that
1. Assigns the value "2" to `a`
2. Assigns the value "4" to `b`
3. Multiplies `a` times `b` and saves it as a variable, `c`.

<!-- grade_id: cell-3d608ef3ac4ec904, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6de2e461718a7ec4, points: 10, grade: true, solution: false -->
```python
## TESTS FOR Q4, worth 10 points [Note: Hidden Tests.]
assert a and b and c

```

## Q4
How are you feeling so far? Create a variable "feeling" and assign your response to it.

- A) I've got this.
- B) This isn't so bad!
- C) I had a lot of trouble getting this to work.
- D) I'm deeply, terribly confused.

<!-- grade_id: cell-9e5ee257e7bd6c04, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6465227207217eba, points: 5, grade: true, solution: false -->
```python
## TESTS FOR Q5, worth 5 points
assert feeling
```