# Assignment 1: Fixing Functions

**PLEASE DO NOT CHANGE THE NAME OF THIS FILE.**

**PLEASE DO NOT COPY & PASTE OR DELETE CELLS INCLUDED IN THE ASSIGNMENT.**


## How to complete assignments

Whenever you see:

```
# YOUR CODE HERE
raise NotImplementedError()
```

You need to **replace (in other words, delete) these lines of code with code that answers the questions** and meets the specified criteria. Make sure you remove the 'raise' line when you do this (or your notebook will raise an error, regardless of any other code, and thus fail the grading tests).

You should write the answer to the questions in those cells (the ones with `# YOUR CODE HERE`), but you can also add extra cells to explore / investigate things if you need / want to. 

Any cell with `assert` statements in it is a test cell. You should not try to change or delete these cells. Note that there might be more than one assert that tests a particular question. 

If a test does fail, reading the error that is printed out should let you know which test failed, which may be useful for fixing it.

Note that some cells, including the test cells, may be read only, which means they won't let you edit them. If you cannot edit a cell - that is normal, and you shouldn't need to edit that cell.

### Tips & Tricks

The following are a couple tips & tricks that may help you if you get stuck on anything.

**Printing Variables**

You can (and should) print and check variables as you go. This allows you to check what values they hold, and fix things if anything unexpected happens.

**Restarting the Kernel**

* If you run cells out of order, you can end up overwriting things in your namespace.
* If things seem to go weird, a good first step is to restart the kernel, which you can do from the kernel menu above.
* Even if everything seems to be working, it's a nice check to 'Restart & Run All', to make sure everything runs properly in order.


# Part I: Formative TV shows

You'd like to be able to predict someone's favorite TV show based on their age. You have a general idea of what TV shows are popular for each age range: 

| Age   | Show                   |
|-------|------------------------|
| 30–39 | Friends                |
| 20–29 | The Office             |
| 10–19 | Pretty Little Liars    |
| 0–9   | Bluey                  |

## Step 1
In order to accomplish this, you'll first need to use the `input()` function. We haven't seen this in class, so first, let's try it out. Use the input function to ask someone "How old are you?" and assign it to `age`. If you're having trouble figure out how input works, ask the internet and/or chatGPT!

<!-- grade_id: cell-f1df5171fb9538d4, points: 5, grade: true, solution: true -->
```python
# your code here
raise NotImplementedError
```

## Step 2
By default, if you give it a whole number, the `input` function doesn't generate a numerical object, it generates a string. In the cell below, cast age into a numerical data type that enables decimal points. Call this variable `num_age`.

<!-- grade_id: cell-672201d4f50cf61c, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-fbe454aacaa8f69e, points: 5, grade: true, solution: false -->
```python
## TESTS FOR Q2, worth 5 points
assert num_age
assert isinstance(num_age,float)
```

## Step 3
Next, we need to create a sequence of if/else statements to determine whether someone is in a given age range. Create a **list** called `solutions` that contains the letters of the following combinations of if and elif that (with the right conditional statements) *could* work to address this task. Each element in your list should be a string containing one letter. For example, if you thought C and D were the only possible solutions, you would write `solutions = ['C','D']`.

- A. Four `if` statements (no `elif`)
- B. Three `if` statements followed by one `elif` statement
- C. One `if` statement, two `elif` statements
- D. Four `elif` statements (no `if`)

<!-- grade_id: cell-cb824717813d567d, points: 0, grade: true, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-f6a143df6105fbb7, points: 2.5, grade: true, solution: false -->
```python
## TESTS FOR Q2, worth 10 points [Note: Includes A Hidden Test]
assert solutions

```

<!-- grade_id: cell-295095a69df0d7b6, points: 5, grade: true, solution: false -->
```python
## HIDDEN TESTS
```

## Step 4
Now that we've determined how to ask someone their age and determine whether it is in a particular range, we can combine this into a function called `predict_show()`. An incorrect version of this function is below. There are **three** errors. Put the complete correct function in your response. The correct function should do two things: ask someone their age and print their favorite show based on the table above.

```
age = float(input("How old are you?"))

def predict_show(age):
    if age > 30 and age < 40:
        show = 'Friends'
    elif age >= 20 and < 30:
        show = 'The Office'
    elif age >= 10 or age <= 19:
        show = 'Pretty Little Liars'
    elif age < 10:
        show = 'Bluey'
        
    print('Your favorite show is',show)
    
    return show
 
```

<!-- grade_id: cell-9e78e70651eb06e0, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-fda93d52cfb41b7e, points: 5, grade: true, solution: false -->
```python
## TESTS FOR Q3, worth 15 points [Note: Includes hidden tests]
assert callable(predict_show)
```

<!-- grade_id: cell-5c81dea147f47859, points: 5, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-9cacf1d210479179, points: 5, grade: true, solution: false -->
```python
## HIDDEN TESTS
```

# Part II: Jumbled potter

For Part II, the following lists are provided for you. These lists contain the stats for various Harry Potter characters, where index 0 is Harry's statistics, index 1 is Hermione's, etc. 

First, make sure you run the following cell to assign these lists.

```python
characters = ['Harry','Hermione','Ron','Neville']
height = [180, 167,190,172]
fears = ['dementor','failure','spider','snape']
```

## Step 1
We, like the rest of Hogwarts, would like to leave poor Neville behind. Either using slicing or a list method, create lists of the **same name** that do *not* include Neville.

<!-- grade_id: cell-df713955a1fe41c6, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-fa359912332aa55c, points: 5, grade: true, solution: false -->
```python
## Worth 5 points [Note: includes hidden tests]
assert isinstance(characters,list)
assert isinstance(height,list)
assert isinstance(fears,list)

```

## Step 2

Next, store the character names and fears into a dictionary called `potter_stats` where each **key** is a characters name and the **value** is their fear. You can do this manually or by using the `zip()` method.

<!-- grade_id: cell-1e532913b17439f9, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-eb74718e856a7b1a, points: 2.5, grade: true, solution: false -->
```python
## Worth 5 points
assert isinstance(potter_stats,dict)
```

<!-- grade_id: cell-809fb6b53997f34d, points: 2.5, grade: true, solution: false -->
```python
## HIDDEN TESTS
```

# Part III: Building an mRNA sequence

In this final part, we'll build an mRNA sequence from scratch. We need the entire sequence to start with our start codon, `AUG`.  After, we'll create a string of nucleotides for **three** Alanine codons (`GCU`) followed by a [stop codon](https://www.genome.gov/genetics-glossary/Stop-Codon), `UAA`.

(For a refresher on how this works, see [this Khan Academy article](https://www.khanacademy.org/science/ap-biology/gene-expression-and-regulation/translation/a/the-genetic-code-discovery-and-properties).)

Build this sequence by following these steps:
1. Create the triple Alanine sequence by multiplying `GCU` time three. Assign this to a variable called `alanine_3`.
2. Build the entire sequence by adding the start codon, `alanine_3`, and your stop codon. Assign this to a variable called `mRNA`.

<!-- grade_id: cell-1b4cb38ac357f6ca, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-f81c66c425da22df, points: 2.5, grade: true, solution: false -->
```python
# Worth 7.5 points total
assert isinstance(alanine_3, str)
assert isinstance(mRNA, str)
```

<!-- grade_id: cell-5e8378bb5be6117b, points: 5, grade: true, solution: false -->
```python
## HIDDEN TESTS
```