# A2: Genes, Neurons, Loops

**PLEASE DO NOT CHANGE THE NAME OF THIS FILE.**

**PLEASE DO NOT COPY & PASTE OR DELETE CELLS INCLUDED IN THE ASSIGNMENT.**


## How to complete assignments

Whenever you see:

```
# YOUR CODE HERE
raise NotImplementedError()
```

You need to **replace (meaning, delete) these lines of code with code that answers the questions** and meets the specified criteria. Make sure you remove the 'raise' line when you do this (or your notebook will raise an error, regardless of any other code, and thus fail the grading tests).

You should write the answer to the questions in those cells (the ones with `# YOUR CODE HERE`), but you can also add extra cells to explore / investigate things if you need / want to. 

Any cell with `assert` statements in it is a test cell. You should not try to change or delete these cells. Note that there might be more than one assert that tests a particular question. 

If a test does fail, reading the error that is printed out should let you know which test failed, which may be useful for fixing it.

Note that some cells, including the test cells, may be read only, which means they won't let you edit them. If you cannot edit a cell - that is normal, and you shouldn't need to edit that cell.


## Tips & Tricks

The following are a couple tips & tricks that may help you if you get stuck on anything.

#### Printing Variables
You can (and should) print and check variables as you go. This allows you to check what values they hold, and fix things if anything unexpected happens.

#### Restarting the Kernel
- If you run cells out of order, you can end up overwriting things in your namespace. 
- If things seem to go weird, a good first step is to restart the kernel, which you can do from the kernel menu above.
- Even if everything seems to be working, it's a nice check to 'Restart & Run All', to make sure everything runs properly in order.

# Part I: Counting Genes
After performing a sequence analysis, you discover that the sequence in one region of interest is the following:
``` GGATAGTGCTTGAGTCTGTGCTGACCTATTATTAAGAACTAAATGGACAATATTATGGAGCATCTCATGTATAAATTGGTGCGTAAAATCGTTGGATCTCTCTTCTAAGTACATCCTACTATAACAATCAAGAAAAACAAGAAAATCGGACAAAACAATCAAGTATGGATTCTAGAACAGTTGGTATATTAGGAGGGGGACAATTGGGACGTATGATTGTTGAGTAAGCTAACAGGCTCAACATTAAGACGGTAATACTAGATGCTGAAAATTCTCCTGCCAAACAAATAAGCAACTCCAATGACCACGTTAATGGCTCCTTGTCCAATCCTCTTGATATCGAAAAACTAGCTGAAAAATGTGATGTGCTAACGATTGAGATTGAGCATGTTGATGTTCCTACACTAAAGAATCTTCAAGTAAAACATCCCAAATTAAAAATCTACCCTTCTCCAGAAACAATCAGATTGATACAAGACA
```
<mark><b>First, run the cell below to assign the sequence to a variable, `sequence`.</b></mark>

```python
sequence = 'GGATAGTGCTTGAGTCTGTGCTGACCTATTATTAAGAACTAAATGGACAATATTATGGAGCATCTCATGTATAAATTGGTGCGTAAAATCGTTGGATCTCTCTTCTAAGTACATCCTACTATAACAATCAAGAAAAACAAGAAAATCGGACAAAACAATCAAGTATGGATTCTAGAACAGTTGGTATATTAGGAGGGGGACAATTGGGACGTATGATTGTTGAGTAAGCTAACAGGCTCAACATTAAGACGGTAATACTAGATGCTGAAAATTCTCCTGCCAAACAAATAAGCAACTCCAATGACCACGTTAATGGCTCCTTGTCCAATCCTCTTGATATCGAAAAACTAGCTGAAAAATGTGATGTGCTAACGATTGAGATTGAGCATGTTGATGTTCCTACACTAAAGAATCTTCAAGTAAAACATCCCAAATTAAAAATCTACCCTTCTCCAGAAACAATCAGATTGATACAAGACA'
```

## Q1
You'd like to know if this region of DNA codes for Phenylalanine, which is coded by either `TTT` or `TTC`. Do the following steps to figure out whether either of these codons are in your data.

1. Create a Boolean `has_TTT` that indicates whether `TTT` is in your DNA string.
2. Create a Boolean `has_TTC` that indicates whether `TTC` is in your DNA string.
3. Create a Boolean `has_phenyl` that is a Boolean containing information about whether or not `TTT` **or** `TTC` is in your sequence.

<!-- grade_id: cell-ce54a1d04d034e5f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c9ebff25941c657d, points: 2, grade: true, solution: false -->
```python
# Tests for Q1. Worth 4 points.
assert isinstance(has_TTT,bool)
assert isinstance(has_TTC,bool)
assert isinstance(has_phenyl,bool)
```

<!-- grade_id: cell-3066b1e70a443d9e, points: 2, grade: true, solution: false -->
```python
# Q1 hidden tests. Worth 2 points.
```

<!-- grade_id: cell-957d006d698c40a4, points: 2, grade: true, solution: false -->
```python
# Q1 hidden tests. Worth 2 points.
```

<!-- grade_id: cell-1a248737dae9fc95, points: 2, grade: true, solution: false -->
```python
# Q1 hidden tests. Worth 2 points.
```

## Q2
Wait, you just realized that you forgot about *codons*! Gene sequences are read by the ribosome in threes. So, we have to search through each third chunk of DNA to see if that is a phenylalanine. 

Use the following steps to set this up:
1. Determine the number of codons in your sequence by dividing the length of the sequence by three. Assign this to `num_codons_float`.
2. Convert `num_codons_float` to an integer, and assign this to `num_codons`. This is how many times you'll need to run your loop in the next question.

<!-- grade_id: cell-350680e16ea610e3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ce657ccd66347d96, points: 2, grade: true, solution: false -->
```python
# Tests for Q2. Worth 2 points.
assert isinstance(num_codons_float,float)
assert isinstance(num_codons,int)
```

<!-- grade_id: cell-17183d7b1e6d04c5, points: 2, grade: true, solution: false -->
```python
# Q2 hidden tests. Worth 2 points.
```

<!-- grade_id: cell-a8da3cbd7ba3f1f8, points: 2, grade: true, solution: false -->
```python
# Q2 hidden tests. Worth 2 points.
```

## Q3

Now that we know how many codons there are, we can use this to determine how many times we need to run our loop.

Create a for loop that looks through each codon and determines whether or not it is a `TTT` or `TTC` sequence by following these steps:
1. Create a for loop that runs through a range equal to the number of codons in your sequence.
2. For each subsequent codon, we'll need to slice the sequence. The first codon is `sequence[0:3]`, the second is `sequence[3:6]`, then `sequence[6:9]`, etc.

    **Hint**: you can programatically identify these start and stop points using a `start_id` that is equal to the codon number (`i`) times 3, and a `stop_id` that is equal to the codon number + 1 times 3. You can then use these ids to slice your sequence differently in each loop, and assign this slice to `this_codon`.
    

3. Within each loop, check whether `this_codon` is equal to `TTT` or `TTC`.
4. Use a counter (`num_phenyl`) to add up how many of these sequences you find. Print this after your loop.

<!-- grade_id: cell-8113842c386da92c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-23bfe0608d6d75ca, points: 2, grade: true, solution: false -->
```python
## Tests for Q3. Worth 2 points.
assert isinstance(num_phenyl,int)
assert isinstance(this_codon,str)
```

<!-- grade_id: cell-868d9cb878d57be6, points: 10, grade: true, solution: false -->
```python
## Hidden tests for Q3. Worth 10 points.
```

As you may remember, RNA has uracil instead of thymine.

## Q4

Convert our DNA sequence to RNA by replacing every "T" with a "U". Assign this to `sequence_RNA`. Don't overthink this -- you may use built-in methods of strings to complete this task!

<!-- grade_id: cell-c6bc0abe7083045b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8b0d6f513e22512d, points: 2, grade: true, solution: false -->
```python
## Tests for Q4. Worth 2 points.
assert isinstance(sequence_RNA,str)
```

<!-- grade_id: cell-75a7be882ba56cfe, points: 2, grade: true, solution: false -->
```python
##  Hidden tests for Q4. Worth 2 points.
```

# Part I: Neuron Spikes

We can think of neuron spike trains as a list of booleans: did the cell spike (fire an action potential), or not?

Let's say that we sample our neuron <font color='red'>**once a millisecond**</font> to see if it spikes (`spike_train`). A **True** value means the neuron spiked in that sample, a **False** value means it did not spike. In other words, `spike_train[0]` represents the first millisecond of the recording, `spike_train[1]` represents the second millisecond, etc. As a further extension of this, the length of `spike_train` is the length of the recording, **in milliseconds**.

## Q5

1. Programmatically count how many times your neuron fires an action potential within the entire recording and assign this to `num_spikes`.
2. Determine the length of your spike train **in milliseconds** and assign this to `train_length`.
3. Determine your neuron's firing rate (# of spikes per second) by dividing the number of spikes by the length of the train, and converting it into seconds from milliseconds. Assign this to `spike_rate`.

<mark><b>Make sure you first run the cell below to assign the spike train.</b></mark>

```python
spike_train = [False,False,True,False,False,False,False,True,False,False,False,False,False,False,False,True,False,True,False,False,False,False,False,True,False,False,True,False,False,False,False]
```

<!-- grade_id: cell-e0770b9af3241d8d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-47e2478ddb011667, points: 2, grade: true, solution: false -->
```python
## Tests for Q5, worth 2 points
assert isinstance(num_spikes,int)
assert isinstance(train_length,int)
assert isinstance(spike_rate,float)
```

<!-- grade_id: cell-800ff2ffd125ad82, points: 3, grade: true, solution: false -->
```python
##  Hidden tests for Q5. Worth 3 points.
```

<!-- grade_id: cell-df8ff83a2b8cc3c2, points: 3, grade: true, solution: false -->
```python
##  Hidden tests for Q5. Worth 3 points.
```

<!-- grade_id: cell-fcd22f86ff038588, points: 3, grade: true, solution: false -->
```python
##  Hidden tests for Q5. Worth 3 points.
```

## Q6
As your final challenge, write a function called `calculate_spike_rate` that takes in `spike_train` and returns the `spike_rate`. First, you'll write the function, and then, we'll use it.

You should use the code you wrote above to complete this task. In the first cell below, write your function. 

<!-- grade_id: cell-2b334fd9bf0d82fb, points: 0, grade: false, solution: true -->
```python
### WRITE YOUR FUNCTION HERE

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1794c374b88210f4, points: 2, grade: true, solution: false -->
```python
## Tests for Q6, worth 3 points. 
# The line below checks that you have a function called "calculate_spike_rate"
"calculate_spike_rate" in dir()
```

Finally, let's test your function on a new spike train. 

<mark><b>Run the cell below to assign `new_spike_train`.</b></mark>

Then, in the following cell, use your function to calculate the spike rate in this new spike train and assign the output to `my_spike_rate`.

```python
new_spike_train = [True,False,False,False,False,False,False,False,False,False,True,False,False,False,False,False,False,False,False,False,True,False,False,False,False,False,False,False,False,False]
```

<!-- grade_id: cell-5ab3a1636aab9d4c, points: 0, grade: false, solution: true -->
```python
## CALL YOUR FUNCTION HERE

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ffd0aa67f2561125, points: 1, grade: true, solution: false -->
```python
## Tests for Q6, worth 1 point.
assert isinstance(my_spike_rate,float)
```

<!-- grade_id: cell-882256abf8e49227, points: 1, grade: true, solution: false -->
```python
## Hidden tests for Q6, worth 1 point.
```

<!-- grade_id: cell-47796b38be433827, points: 4, grade: true, solution: false -->
```python
## Hidden tests for Q6, worth 4 points.
```