# Assignment a3

## Q1

Write a `while` loop that takes a number and continues to subtract 7 from it, until `number` is **negative** (note that 0 is not negative). The loop should start with `number = 21` (provided for you) and include a `counter` which counts how many times the loop executes.

<!-- grade_id: cell-98a8ccfa9b2e1fce, points: 0, grade: false, solution: true -->
```python
number = 21

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-786293d03fb8406b, points: 3, grade: true, solution: false -->
```python
# Tests for Q1, worth 11 points. Note: hidden tests below.
assert isinstance(number, int)
assert isinstance(counter, int)
```

<!-- grade_id: cell-4ef1ad8ddffdf312, points: 4, grade: true, solution: false -->
```python
# First hidden test cell for Q1
```

<!-- grade_id: cell-b194e3d47aa026f1, points: 4, grade: true, solution: false -->
```python
# Second hidden test cell for Q1
```

## Q2
In this question, we'll deepen our understanding of mutability in lists. 

1. Create a list of three animals (**bear, monkey, zebra**) and assign it to `zoo`.
2. Assign `zoo` to `zoo_alias`.
3. Change the first item in `zoo` to **tiger**.
4. Change the third item in `zoo_alias` to **lion**.
5. Use a membership operator to test whether **lion** is in `zoo`. Assign this as a boolean to `has_lions`. 

<!-- grade_id: cell-ed5b1ad32cd07039, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8678b95de5e898f0, points: 3, grade: true, solution: false -->
```python
# Tests for Q2, worth 11 points. Note: hidden tests.
assert isinstance(zoo, list)
assert isinstance(zoo_alias, list)
assert isinstance(has_lions,bool)
```

<!-- grade_id: cell-3554c5a0a624cdbd, points: 4, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-ced1b5006860c6b2, points: 4, grade: true, solution: false -->
```python
# Second hidden test cell for Q2
```

## Q3
### Zoo inventory

Given the zoo composition below (be sure to run the cell below), write a code snippet that creates a dictionary containing each type of animal and how many there are. In other words, the keys in the dictionary should be the animals, and the values should be the animal count. Call the dictionary `zoo_counts`.

**Hints**:
* You can create an empty dictionary using empty brackets.
* You can check if a dictionary already has a key using a membership operator *or* you can find unique values in a list by converting it to an array and using `np.unique()`.
* You can refer to a specific value in the dictionary using the syntax `dictionary_name[key_name]`. For example, if the dictionary is called `zoo_counts` and you are trying to access the value assigned to tiger, you would use the notation `zoo_counts['tiger']`.

```python
# Run this cell to establish the zoo
sandiego_zoo = ['tiger','monkey','lion','lion','bear','monkey','monkey','monkey','tiger','monkey','monkey']
```

<!-- grade_id: cell-2e43c6efad974686, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-06e0be1892515253, points: 2, grade: true, solution: false -->
```python
# Tests for Q3, worth 9 points. Note: hidden tests.
assert isinstance(zoo_counts, dict)
```

<!-- grade_id: cell-a81a77e1fc39c625, points: 3, grade: true, solution: false -->
```python
# Q3 hidden test cell #1
```

<!-- grade_id: cell-ab2e051b05ab4142, points: 4, grade: true, solution: false -->
```python
# Q3 hidden test cell #2
```

## Q4a

### Create a burrito class

Create a class called `Burrito`. This `Burrito` class should have:

1. An instance-specific attribute, (`contents`) that is a list of contents in the burrito. For example, when you create a new burrito, you will need to provide a list of contents:
```
  my_burrito = Burrito(['lettuce','fries','steak'])
```
and this will update the `contents` attribute.

2. A second instance-specific attribute (`is_california`) that is a **boolean** that indicates if the burrito is a california burrito (in other words, if it has **fries** in it!).
3. A method, (`add_fries`). If there were not already fries in the contents list, this method should add **fries** to the list of `contents` *and* update `is_california` to be true.

<!-- grade_id: cell-5a64823b63dd8ffc, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-acdd925ced2eeb3b, points: 3, grade: true, solution: false -->
```python
# Tests for Q5, worth 6 points.
import inspect
assert inspect.isclass(Burrito) #Checks if burrito is a class
```

<!-- grade_id: cell-7367e30633361953, points: 3, grade: true, solution: false -->
```python
# Second (hidden) test cell for Q5
```

## Q4b
### Create an instance of your burrito class

Below, create an instance of your burrito class called `favorite_burrito` with your **three** favorite ingredients (anything will work!)

<!-- grade_id: cell-026003a769ec5776, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-fb0ec96a89994593, points: 3, grade: true, solution: false -->
```python
# Tests for Q5b, worth 7 points. Note: hidden tests.
assert len(favorite_burrito.contents) == 3
```

<!-- grade_id: cell-a7773ff181fa8041, points: 4, grade: true, solution: false -->
```python
# Q5b hidden test cell
```

## Q4c
### Test `add_fries` method.
Test your add fries method below.

<!-- grade_id: cell-78cefe7646d503ad, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1aa9b59a88a6d2b2, points: 2, grade: true, solution: false -->
```python
# Tests for Q5c, worth 6 points.
assert 'fries' in favorite_burrito.contents
```

<!-- grade_id: cell-4e7d68e05cc8928b, points: 4, grade: true, solution: false -->
```python
# Second (hidden) test cell
```