# Assignment a3

## Q1

Write a `while` loop that takes a number and continues to subtract 7 from it, until `number` is **negative** (note that 0 is not negative). The loop should start with `number = 21` (provided for you) and include a `counter` which counts how many times the loop executes.

<!-- grade_id: cell-98a8ccfa9b2e1fce, points: 0, grade: false, solution: true -->
```python
number = 21

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-786293d03fb8406b, points: 5, grade: true, solution: false -->
```python
# Tests for Q1, worth 2 points.
assert isinstance(number, int)
assert isinstance(counter, int)
```

<!-- grade_id: cell-10c81f3ad5dad643, points: 2.5, grade: true, solution: false -->
```python
# Tests for Q1, worth 2.5 points.
```

<!-- grade_id: cell-b194e3d47aa026f1, points: 2.5, grade: true, solution: false -->
```python
# Second hidden test cell for Q1, worth 2.5 points.
```

## Q2
In this question, we'll deepen our understanding of mutability in lists. 

1. Create a list of three animals (**bear, monkey, zebra**) and assign it to `zoo`.
2. Assign `zoo` to `zoo_alias`.
3. Change the first item in `zoo` to **tiger**.
4. Change the third item in `zoo_alias` to **lion**.
5. Use a membership operator to test whether **lion** is in `zoo`. Assign this as a boolean to `has_lions`. 

<!-- grade_id: cell-ed5b1ad32cd07039, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8678b95de5e898f0, points: 2, grade: true, solution: false -->
```python
# Tests for Q2, worth 2 points

assert isinstance(zoo, list)
assert isinstance(zoo_alias, list)
assert isinstance(has_lions,bool)
```

<!-- grade_id: cell-6108b1a6f08829ca, points: 2, grade: true, solution: false -->
```python
# Hidden tests for Q2, worth 2 points.
```

<!-- grade_id: cell-ced1b5006860c6b2, points: 2, grade: true, solution: false -->
```python
# Second hidden test cell for Q2, worth 2 points.
```

## Q3
### Zoo inventory

Given the zoo composition below (be sure to run the cell below), write a code snippet that creates a dictionary containing each type of animal and how many there are. In other words, the keys in the dictionary should be the animals, and the values should be the animal count. Call the dictionary `zoo_counts`.

**Hints**:
* You can create an empty dictionary using empty brackets.
* You can check if a dictionary already has a key using a membership operator *or* you can find unique values in a list by converting it to an array and using `np.unique()`.

```python
# Run this cell to establish the zoo
sandiego_zoo = ['tiger','monkey','lion','lion','bear','monkey','monkey','monkey','tiger','monkey','monkey']
```

<!-- grade_id: cell-2e43c6efad974686, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-06e0be1892515253, points: 2, grade: true, solution: false -->
```python
# Tests for Q3, worth 2 points. 

assert isinstance(zoo_counts, dict)
```

<!-- grade_id: cell-ead8e0cddd067393, points: 2, grade: true, solution: false -->
```python
# Hidden tests for Q3, worth 2 points.
```

<!-- grade_id: cell-ab2e051b05ab4142, points: 2, grade: true, solution: false -->
```python
# Q3 test cell #2, worth 2 points.
```

## Q4
### Are we there yet?
Using `input()` and a `while` loop, write a function called `backseat_banter` that checks whether the car is there yet. If the `response` is yes (**in ANY form of capitalization**, e.g. YES, yes, Yes), exit the function and print **We are finally there!**. If the answer is anything but **Yes**, ask again. 

<!-- grade_id: cell-dd3585915c7dd609, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9b571fde05195eb0, points: 10, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Q5a

### Create a burrito class

Create a class called Burrito. This burrito should have:
* An attribute (`is_california`) that is a **boolean** that indicates if the burrito is a california burrito (in other words, if it has **fries** in it!).
* An instance-specific attribute, (`contents`) that is a list of contents in the burrito.
* A method, (`add_fries`). If there were not already fries in the contents list, this method should add **fries** to the list of `contents` *and* update `is_california` to be true.

<!-- grade_id: cell-5a64823b63dd8ffc, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-acdd925ced2eeb3b, points: 2, grade: true, solution: false -->
```python
# Tests for Q5, worth 2 points.
import inspect
assert inspect.isclass(Burrito) #Checks if burrito is a class
```

<!-- grade_id: cell-7367e30633361953, points: 2, grade: true, solution: false -->
```python
# Second test cell for Q5, worth 2 points
assert inspect.isfunction(Burrito.add_fries) #Checks if add_fries is a method
```

## Q5b
### Create an instance of your burrito class

Below, create an instance of your burrito class called `favorite_burrito` with your three favorite ingredients (anything will work!)

<!-- grade_id: cell-026003a769ec5776, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-fb0ec96a89994593, points: 2, grade: true, solution: false -->
```python
# Tests for Q5b, worth 2 points.
assert len(favorite_burrito.contents) == 3
```

<!-- grade_id: cell-a7773ff181fa8041, points: 4, grade: true, solution: false -->
```python
# Q5b hidden test cell, worth 4 points.
```

## Q5c
### Test `add_fries` method.
Test your add fries method below.

<!-- grade_id: cell-78cefe7646d503ad, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1aa9b59a88a6d2b2, points: 8, grade: true, solution: false -->
```python
# Tests for Q5c, worth 9 points.
assert 'fries' in favorite_burrito.contents
assert favorite_burrito.is_california == True 
```