# Assignment a4
In this assignment, you'll work through a dataset that is similar in format to our inflammation dataset. 

## Step 1: Import the necessary packages

Below, import `numpy` as `np` and `matplotlib.pyplot` as `plt`.

<!-- grade_id: cell-fb75c1baaa034a79, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a663c96201b646a1, points: 5, grade: true, solution: false -->
```python
# Tests for Step 1, 5 pts
import sys
assert 'numpy' in sys.modules
assert 'matplotlib.pyplot' in sys.modules
```

<!-- grade_id: cell-9d36a0f8e7172bab, points: 5, grade: true, solution: false -->
```python
# Two more tests for Step 1, 5 pts
assert 'np' in dir()
assert 'plt' in dir()
```

## Step 2: Load our data as a numpy array
In the cell below, load our file, called 'pain_data.csv'.

1. Assign this array to a variable called `pain_data`.
2. Assign the shape of this array to `data_shape`.

<!-- grade_id: cell-e89d8ae3668d086e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-199dc904caa4e627, points: 5, grade: true, solution: false -->
```python
# Tests for Step 2, 5 pts
assert isinstance(pain_data,np.ndarray)
assert isinstance(data_shape,tuple)
```

<!-- grade_id: cell-3bc2d88f15ed1161, points: 5, grade: true, solution: false -->
```python
# Hidden tests for Step 2, 5 pts
```

## Step 3: Divide into control and treatment groups
In this data set, each row is a patient, and each column is the hour of the test. The first 60 patients (in the first 60 rows) in this study were in the control group. The second group of patients were in a treatment group. Divide `pain_data` into two separate `control` and `treatment` arrays accordingly.

<!-- grade_id: cell-7321442fff968ae4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-252c7126f09f9113, points: 5, grade: true, solution: false -->
```python
# Tests for Step 3, worth 5 pts
assert isinstance(control,np.ndarray)
assert isinstance(treatment,np.ndarray)
```

<!-- grade_id: cell-82b19a656d13da5e, points: 5, grade: true, solution: false -->
```python
# Hidden tests for Step 3, worth 5 pts
```

## Step 4: Create a histogram to compare baseline measures
As a first check, we want to make sure that our control and treatment groups have similar baseline distributions. **The first hour of this dataset (the first column) has the baseline values.**

To compare our baseline values, generate two side-by-side subplots. The left plot should show a histogram (using `plt.hist()`) of the control group values, the right plot should show a histogram of the treatment group values. **You do not need to add any axis labels or titles to these plots**. 

Please be sure that you have a `plt.show()` command at the end of your cell.

*Note*: This question is manually graded. There is no test cell.

<!-- grade_id: cell-c0f932bb6f655aae, points: 10, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Step 5: Generate an average trajectory for each group
Next, we'd like to generate an average pain calculation across the entire eight hour period for our control and treatment groups. In other words, what is the average pain trajectory from hours 0 to 7 for each of our groups?

Using `np.mean()`, calculate an average value across all control patients and assign it to `control_mean`. Calculate an average value across all treatment patients and assign it to `treatment_mean`. 

<!-- grade_id: cell-c553db54a5548b20, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3e9feb695c4a99ad, points: 2, grade: true, solution: false -->
```python
# Tests for step 5, 2 pts
assert isinstance(control_mean,np.ndarray)
assert isinstance(treatment_mean,np.ndarray)
```

<!-- grade_id: cell-1af01308c5e29679, points: 4, grade: true, solution: false -->
```python
# Hidden tests for step 5, 4 pts
```

<!-- grade_id: cell-19e3b49cf403eb42, points: 4, grade: true, solution: false -->
```python
# Hidden tests, part II, 4 pts
```

## Step 6: Plot the average trajectory

Below, create **one** plot that contains all eight hours of data for your control and treatment groups, plotted as two separate lines.

### Tips:
* Write two `plt.plot()` lines after one another to draw multiple lines on the same plot. For example:
```plt.plot(x)
plt.plot(y)
```
* Create a legend with `plt.legend()`. This function takes one argument that is a list of all of the legend labels. For example:
```
plt.legend(['x data','y data'])
```

### Grading:
For full credit on this manually graded question, your graph must include:
* Two separate lines on one plot that represent the mean pain across eight hours of data for control and treatment groups.
* A legend that shows which group is which
* An x-axis label `Hour`
* A y-axis label `Pain`
* A final line with a `plt.show()` command

<!-- grade_id: cell-b05ab460594cba31, points: 15, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Step 7: Check to see how many patients did not respond to treatment
Finally, let's see if there are any patients that did *not* respond to treatment. We will define patients as resistant if their pain rating was 10 or higher throughout the treatment period.

Write code below that determines how many of the treatment patients *never* had a pain rating lower than 10 (in other words, their lowest pain was always more than 9) and assigns this to `resistant_patients`.

### Tips:
* You can calculate the minimum values using `np.min()`.
* One way to do this includes a `for` loop that loops through each patient, counting up the `resistant_patients` along the way.

<!-- grade_id: cell-e826f777bc010a97, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b4521bcecb15a52d, points: 5, grade: true, solution: false -->
```python
# Test cell for Step 7 (5 pts)
assert isinstance(resistant_patients,int)
```

<!-- grade_id: cell-919a7c0f516f4aa8, points: 5, grade: true, solution: false -->
```python
# Hidden tests for Step 7 (5 pts)
# Hidden tests
```