# Assignment a6: Linear Algebra


## Q1
Using NumPy, create a two-dimensional vector, where the x-coordinate is 12, and the y-coordinate is 5. Assign this vector to an object called `vector`.

<!-- grade_id: cell-9be52e26d65137a0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f2669f2b98c712d9, points: 2, grade: true, solution: false -->
```python
assert isinstance(vector,np.ndarray)
```

<!-- grade_id: cell-c8a5d94c248de7d0, points: 3, grade: true, solution: false -->
```python
# Hidden tests
```

## Q2

Find the unit vector of the `vector` above and assign it to `vector_unit`.

<!-- grade_id: cell-373893a38dd743ec, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-aa6049e77a396c93, points: 2, grade: true, solution: false -->
```python
assert isinstance(vector_unit,np.ndarray)
```

<!-- grade_id: cell-5fec3fe69690f33b, points: 3, grade: true, solution: false -->
```python
# Hidden tests
```

## Q3

**True or False**: a vector with coordinates (1,1) is a unit vector. Your response should be a Boolean assigned to `Q3`.

<!-- grade_id: cell-100fe29dbe7b3ef3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b19483f3a55c6fad, points: 3, grade: true, solution: false -->
```python
assert isinstance(Q3,bool)

# Hidden tests
```

## Q4

Write a function `vector_addition` that takes two vectors represented as lists and returns their sum. For example, given vectors [1, 2, 3] and [4, 5, 6], the function should return [5, 7, 9].

<!-- grade_id: cell-14d487e20ea1b49a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b3a1ee5263ffacd5, points: 10, grade: true, solution: false -->
```python
# Hidden tests
```

## Q5

Write a function `create_vector` that can take a vector and designed magnitude (length) as inputs, and will return a vector in the same direction, but with a magnitude corresponding to the second input.

<!-- grade_id: cell-491e5108f687aae9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-57544efb48405a39, points: 5, grade: true, solution: false -->
```python
# Hidden tests
```

## Q6

**Enzyme Kinetics: Deciphering Substrate Concentrations**

Consider an enzymatic reaction involving 2 enzymes, denoted as E<sub>1</sub> and E<sub>2</sub>, which are part of a biochemical pathway. These enzymes are influenced by the concentrations of 2 substrates, S<sub>1</sub> and S<sub>2</sub>. Let $c_{E_1}$ and $(c_{E_2}$ represent the catalytic activities of enzymes $E_1$ and $E_2$, respectively. The concentrations of substrates are denoted as $s_1$ and $s_2$.

The relationship between substrate concentrations and enzyme activities is modeled by the following system of linear equations:

\begin{align*}
c_{E_1} &= 2 \cdot s_1 + 3 \cdot s_2 \\
c_{E_2} &= -s_1 + 4 \cdot s_2
\end{align*}

Assume that you have only measured the activities of enzymes ($c_{E_1}=10$ and $c_{E_2}=5$) and are attempting to deduce the concentrations of substrates.

### Step 1
Write the matrix equation representing the system of linear equations in terms of matrices and vectors.

- Create the coefficients matrix and assign it to `coeff_matrix`.
- Assign the inverse of that matrix to `inverse_matrix`.
- Create a vector of the measured enzyme activities and assign it to `enzyme_activities`.

<!-- grade_id: cell-a48a9616b9393ee7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-7dee79db9186a222, points: 2, grade: true, solution: false -->
```python
assert isinstance(coeff_matrix,np.ndarray)
assert isinstance(inverse_matrix,np.ndarray)
assert isinstance(enzyme_activities,np.ndarray)
```

<!-- grade_id: cell-bfdc3de4680312ee, points: 2, grade: true, solution: false -->
```python
# Hidden tests
```

<!-- grade_id: cell-b4301670a154973d, points: 2, grade: true, solution: false -->
```python
# Hidden tests
```

<!-- grade_id: cell-fa542689b9c44e61, points: 2, grade: true, solution: false -->
```python
# Hidden tests
```

### Step 2
Solve the matrix equation to find the concentrations of substrates ($(s_1$ and $s_2$) given the observed enzyme activities ($c_{E_1}$ and $c_{E_2}$). Assign the output of this to `substrate_concentrations`.

<!-- grade_id: cell-1cea2da5c6a71b77, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-092c49f6e30ae5b0, points: 2, grade: true, solution: false -->
```python
assert isinstance(substrate_concentrations,np.ndarray)
```

<!-- grade_id: cell-8e8db7f98fc52fe9, points: 3, grade: true, solution: false -->
```python
# Hidden tests
```

## Q7
One of the biggest use cases of linear algebra in biology is to enable dimensionality reduction. Below, we'll create a dataset containing 100 cells, with 8 different types of information about each cell (stored in `cell_fields`). We have already **mean centered** this dataset, so it is ready for analysis.

**Run the cell below to store the dataset and preview it.**

```python
cell_fields = ['latency', 'volume', 'depth', 'maxrate', 'spont', 'width', 'axon', 'dendrite']
x_data = pd.read_csv('neurons.csv')
x_data = x_data[cell_fields] # only use cell_fields columns
x_data.head()
```

The first step is to create the covariance matrix. This will show us how each of the measurements of our cell are related. **Run the cell below to create the covariance matrix and visualize it.**

```python
import matplotlib.pyplot as plt

covariance_matrix = x_data.cov()

plt.imshow(covariance_matrix, cmap='viridis')
plt.colorbar()
plt.title('Covariance Matrix (\u03A3)')
plt.xticks(range(len(cell_fields)), cell_fields, rotation='vertical')
plt.yticks(range(len(cell_fields)), cell_fields)

plt.show()
```

Now we are ready to compute the eigenvectors and eigenvalues and then visualize our results!

### Steps 1-3
1. Compute the eigenvectors and eigenvalues using `np.linalg.eigh()` and assign the outputs to `eigenvalues` and `eigenvectors`. [Documentation for `np.linalg.eigh()` is here.](https://numpy.org/doc/stable/reference/generated/numpy.linalg.eigh.html)
2. `eigenvectors` should be an by an 8 array, where each column is an eigenvector. Slice this array to get the first two eigenvectors (in other words, the first two columns of the data and assign this to `top_eigenvectors`.
3. Compute the dot product between `x_data` and `top_eigenvectors` and assign this to `transformed_data`.

<!-- grade_id: cell-106ffb2ddf2b9ee6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-546e292fa24ccb2f, points: 2, grade: true, solution: false -->
```python
assert isinstance(eigenvalues,np.ndarray)
assert isinstance(eigenvectors,np.ndarray)
assert isinstance(top_eigenvectors,np.ndarray)
```

<!-- grade_id: cell-8f4b25a183e335e2, points: 2, grade: true, solution: false -->
```python

```

### Step 4
Finally, visualize your results by plotting the first column of `transformed_data` on the x-axis of a scatter plot, with the second column on the y-axis. You should see three fairly discrete clumps of datapoints when you do this.

**Note**: You do not need any axis labels on this plot. This step is manually graded.

<!-- grade_id: cell-e4b2b0fa54c63b0a, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```