# Signal Processing

In this assignment, you'll design a data workflow to detect peaks in a biological signal. Along the way, you'll write two functions and include these in a custom "spikes" module.

### About Spike Sorting
The goal here is to take a continuous recording of voltage from a brain and pick out single action potentials (*spikes*) so that we can identify which neurons fired which spikes. You can see this in the bottom panel of the image below (from <a href="https://www.sciencedirect.com/science/article/pii/S0361923015000684">this paper</a>), where a threshold (red line) has been used to identify where the signal crosses a certain voltage. Then, future steps can be taken (not discussed in this assignment) to determine whether spikes come from the same neuron or different neurons. In this assignment, you'll do the first few steps of this process: cleaning and thresholding the data to detect spikes.

![](https://ars.els-cdn.com/content/image/1-s2.0-S0361923015000684-gr1.jpg)


<hr>

### Step 1. Import data

The data is in your assignment folder as `recording.txt`. This file contains a <mark>**one-dimensional recording of voltage (mV) over time. In other words, each entry in the file is a voltage measured at some time point.**</mark>

Import `recording.txt` as a numpy array below, assigned to a variable `recording`. Be sure to import whatever packages you need first.

<!-- grade_id: cell-e29e30b1fac9a023, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-aad92426cc7a8f1c, points: 2, grade: true, solution: false -->
```python
# Tests for Q1, worth 2 points.
assert isinstance(recording, np.ndarray)
```

<!-- grade_id: cell-abbf7bf587d50855, points: 2, grade: true, solution: false -->
```python
# Hidden tests for Q1, , worth 2 points.
```

### Step 2. Check to see if the data needs to be cleaned.
Our equipment has one known failure -- if it misses a sample, it assigns the sample the value -1000 and then continues recording. So, if our data failed the check above, it probably is due to this known issue.

It's tough to know how to fix such outliers (ideally this wouldn't happen!) but one method that isn't *too* egregious is to fill in these outliers with the median value in the dataset. 

To clean our data, write a function `interpolate_outliers` that does the following:

1. Takes in a time series (e.g., the recording we loaded above) as an input
2. Determines the median of the data and assigns this to `recording_median`
3. Determines whether there are any values in the recording that less than -999
4. If it finds those values, it reassigns them to the median of the data.
5. Returns how many values were replaced in an integer (`num_replaced`), a list that contains the indices of those values replaced (`indices_replaced`), and the new recording (`recording_clean`).

#### Hints:
- You can use `enumerate` if you want to loop through something and obtain both the value and index while looping.
- You can return multiple values from a function by using `return a,b,c` and then `a,b,c = function(x)`.
- <mark>**You may want to modify a copy of the data, so that the original data is still stored in a variable. To make a copy of an array, use the syntax `array.copy()`**</mark>

<!-- grade_id: cell-692391d622703879, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0c25b86c484d968a, points: 2, grade: true, solution: false -->
```python
# Tests for Step 2, worth 2 points
import inspect
assert inspect.isfunction(interpolate_outliers)
```

<!-- grade_id: cell-cb9b68de4d2e5228, points: 5, grade: true, solution: false -->
```python
# Hidden tests for Step 2, worth 5 points
```

### Step 3: Generate time stamps
This data was collected at 24000 Hz and is 43 seconds long. <mark>**In other words, 24000 samples were measured per second, and one sample was measured every 1/24000 seconds.**</mark>
    
    
Using `np.arange()`:
1. Generate a vector of timestamps with that many samples with the correct spacing between samples. Assign this vector to `timestamps`.
2. Generate a vector of sample ids (`sample_ids`) that is the same length of recordings, with a step of 1.

<mark>**Hint**: Look back at our signal processing notebook for how to do this step!</mark>

<!-- grade_id: cell-c3a8e7a6b80a6533, points: 0, grade: false, solution: true -->
```python
sampling_freq = 24000

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-82e1485082e0dfe6, points: 2, grade: true, solution: false -->
```python
# Tests for Step 3, worth 2 points
assert isinstance(sample_ids,np.ndarray)
assert isinstance(timestamps,np.ndarray)
```

<!-- grade_id: cell-76d68b3bca2d7e55, points: 2, grade: true, solution: false -->
```python
# Hidden tests for Step 3, worth 2 points
```

<!-- grade_id: cell-8446d37f06c2c413, points: 2, grade: true, solution: false -->
```python
# Hidden tests for Step 3, worth 2 points
```

### Step 4: Visually inspect the data

Plot the first second of your cleaned recording below. Remember that the data was sampled at 24000, so 24000 samples is equivalent to one second of the recording.
* For full credit, label your x and y axes. 

**Note**: This question is manually graded. There is no test cell.

<mark>**To receive full points on this question, you need to add `plt.show()` at the end of your cell.**</mark>

<!-- grade_id: cell-b9cc94f3d3a8ee2a, points: 10, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Step 5. Identify places in the data that cross a threshold
As you can see in the recording above, sometimes the data "spikes" both positively and negatively. These are action potentials! <mark>Our goal is to find these action potentials so that we can more closely inspect their shape.</mark> As a first step, we will determine the times that the data dips **below** a given threshold.

To do so:
1. First create a list of booleans (`threshold_bool`) that is the same length of your recording and is `True` if the recording is less than a variable `threshold` at each sample. **Hint**: You can do this for an entire array by using `recording < threshold`.
2. <mark>**Recreate `sample_ids` in the same way you did in step 3 above.**</mark>
3. Subset your `sample_ids` array so that you are left **only** with the sample IDs where there are threshold crossings, <mark>**and assign this to `crossings`.** Remember that you can subset an array with a boolean expression.</mark>
4. Wrap this in a function called `find_crossings` that takes `recording` and `threshold` as inputs and returns a numpy array called `crossings`.

<mark>**Here is what your resulting function will look like:**</mark>
```
    def find_crossings(recording,threshold):
            threshold_bool = ... # Step 1
            sample_ids = ...     # Step 2
            crossings = ...      # Step 3
    
            return crossings
```


<!-- grade_id: cell-96cb2200d86c76e7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c929cfd41c699819, points: 2, grade: true, solution: false -->
```python
# Tests for Step 5, worth 2 points
assert inspect.isfunction(find_crossings)
```

<!-- grade_id: cell-d4a538ea51936ab0, points: 2, grade: true, solution: false -->
```python
# Hidden tests for Step 5, worth 2 points
```

### Step 6. Add your functions to our `spike` module.
In the same folder as this assignment, you'll see there's a folder called `spike_module`.

Within this folder, there is a file called `spikes.py`. If you open that file, you'll notice that it contains several function definitions. Add the two functions you created above (`find_crossings`,`interpolate_outliers`) at the end of this file by copying and pasting them into spikes.py and saving it.

*Then*, import this new custom module by typing `from spike_module import spikes` below. 

#### Tips
* **If you make changes to the module, you need to re-import it.**
* If you still do not pass the `assert` below after saving changes to `spikes.py`, restart the kernel.

<!-- grade_id: cell-a67071499b6f6787, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-12c98fbf846ccca5, points: 4, grade: true, solution: false -->
```python
## Tests for Step 6, worth 4 points
assert inspect.isfunction(spikes.interpolate_outliers)
assert inspect.isfunction(spikes.find_crossings)
```

### Step 7.  Use your new module to plot spikes!

In the box below:
1. First use your `find_crossings` function to identify all threshold crossings **in your clean recording with a -30 threshold** and assign this to `crossings`.
2. Use the `align_to_minimum` function in the `spikes` module to identify the spikes. Use 0.002 as the search range. Return this as `spks`. **Hint**: To use functions from the spike module, you need to call them with `spikes.<function name>`. For example, `spikes.align_to_minimum`.
3. Use the `extract_waveforms` function to return `waveforms`. Note that `extract_waveforms` takes three arguments: `recording_clean, sampling_freq, spks`.
4. Use `plot_waveforms` to plot. Check the function definition to see which arguments it needs.

<mark>**Hint**: Take a look at `spikes.py` for documentation on what arguments each function
takes.</mark>

<!-- grade_id: cell-e7defb657f0ff017, points: 10, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-86ff332aa8b22a94, points: 2, grade: true, solution: false -->
```python
# Tests for Step 7, worth 2 points.
assert isinstance(crossings,np.ndarray)


```

<!-- grade_id: cell-4a87751c2cdaf44c, points: 3, grade: true, solution: false -->
```python
# Hidden Tests for Step 7, worth 3 points.
```