# Informatics
Our final assignment is focused on the content from Week 9. Completing this assignment will be greatly aided by the Neuroinformatics lecture slides as well as the Jupyter Notebook we worked through in class.

## Q1

What is the name for "The field of computational science that has to do with the analysis of sequences of biological molecules."? Write it as a string and assign it to `Q1_answer` below.

<!-- grade_id: cell-dca2b40e11753d34, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-30e87cc24c8a1796, points: 2.5, grade: true, solution: false -->
```python
# Tests for Q1, including hidden tests (2.5 pts)
assert isinstance(Q1_answer,str)
```

## Q2
What is an API?

* A - An application that acts as a messenger between two software components to communicate and takes requests, translates, and returns responses
* B - A way to organize genomics data
* C - An 'application performance index' which tells you the speed of a module
* D - An advanced way to package Python modules

Assign your answer to `Q2_answer` below. For example, if you wanted to answer `D` you should write `Q2_answer = 'D'`.

<!-- grade_id: cell-fedc47a7e3add3e8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f9c8dbd9657fcef6, points: 2.5, grade: true, solution: false -->
```python
# Tests for Q2, including hidden tests (2.5 pts)
assert isinstance(Q2_answer,str)
```

## Q3
Import the `Entrez` package from Bio (Biopython) below. **Hint**: We did this after installing our packages in the neuroinformatics notebook.

<mark>**Note**: The cell below *will not* work unless you have already run `!pip install Bio --user` to install the Bio module. We did this during lecture.</mark>

<!-- grade_id: cell-56693ce2aa3e3e63, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c9d787339d4342f4, points: 5, grade: true, solution: false -->
```python
## Tests for Q3, 5 pts
import inspect
assert inspect.ismodule(Entrez)
```

## Q4
1. Use `einfo` to generate a list of databases accessible in Entrez. Assign this to handle.
2. Use the `read()` method of Entrez on the `handle` and assign this output to `databases`.
3. Assign the list *within* the `databases` dictionary (the value to the `DbList` key) to `database_list`. In the end, `database_list` should be a list of databases, which are each strings. **Hint**: Inspect the `databases` object. What is the first key in the dictionary?

<!-- grade_id: cell-7400f7930ecc79ed, points: 0, grade: false, solution: true -->
```python
Entrez.email = ... # Optional step to add your email

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4a70b35390bc8f49, points: 2.5, grade: true, solution: false -->
```python
# Tests for Q4 (2.5 pts)
assert isinstance(databases,dict)
assert len(databases) == 1
```

<!-- grade_id: cell-006095b2b99d6d07, points: 2.5, grade: true, solution: false -->
```python
# Hidden tests for Q4 (2.5 pts)
```

## Q5
For **each** of the databases in `database_list`, pull out the description of the database. 

In order to so, follow these steps:
1. Initialize an empty dictionary called `db_dictionary`.
2. Set up a for loop that iterates over each database in `database_list`.
3. Within the loop, generate `handle` which is the output of `Entrez.einfo()` for the database in that loop iteration. In other words, you can use `Entrez.einfo(db=db)`, since `db` will update on each loop.
4. Read the handle using `Entrez.read()` and assign that to a variable called `info`.
5. Get the dictionary assigned to the `DbInfo` key from within `info` and assign that to `this_db`.
6. Assign the `Description` within `this_db` to a new entry in your empty dictionary, with the name of the database as the key. For example, `db_dictionary[db] = this_db['Description']`. <mark>**At the end of this loop, db_dictionary should be a dictionary where each key is the name of a database, and the value is a description of the database.**</mark>

Here is a skeleton of the code:

``` 
db_dictionary = {}        # Step 1

for db in database_list:  # Step 2
    handle = ...          # Step 3
    info = ...            # Step 4
    this_db = ...         # Step 5
    db_dictionary[db] = this_db['Description']
```


<!-- grade_id: cell-4ad5db994294fa0d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3f0d78905a58d89e, points: 10, grade: true, solution: false -->
```python
# Tests for Q5 (10 pts)
assert isinstance(db_dictionary,dict)
```

<!-- grade_id: cell-5f002d76d6124251, points: 5, grade: true, solution: false -->
```python
# Hidden tests for Q5 (5 pts)
```

## Q6
1. Use `esearch` to look for papers in pubmed that mention 'cardiovascular disease' and assign this handle to `pub_handle`. Return the first 100 results, sorted by relevance. [The esearch help page is here](https://www.ncbi.nlm.nih.gov/books/NBK25499/#chapter4.ESearch).
2. Using the `read()` method, assign the output of this search to `pub_record`.
3. Assign the IDs for these (contained in `IdList`) to `paper_ids`.

<!-- grade_id: cell-0c77abedc8548d4f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-88b7b7074ad7f4cb, points: 5, grade: true, solution: false -->
```python
# Tests for Q6 (5 pts)
assert isinstance(paper_ids,list)
assert len(paper_ids) == 100
```

<!-- grade_id: cell-093977b0018aa6c4, points: 5, grade: true, solution: false -->
```python
# Hidden tests for Q6 (5 pts)
```

## Q7
1. Use `efetch` to get the abstract for the **very first paper** in your `paper_ids` above and assign it to `abstract_handle`. Hint: you need to give `efetch` three arguments: `db`, `id`, and `rettype`.
2. Use `read()` on your handle and assign it to `abstract`.

Abstract should be an bytes `xml` (kind of messy) object that contains the abstract, along with lots of other information. If you are getting a "Bio.Entrez.Parser.DictionaryElement" instead of a bytes object, make sure you are using the syntax `handle.read()` and try again. 

<!-- grade_id: cell-56a3fe7e38bfa57b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-fcd671dad0d6b869, points: 2.5, grade: true, solution: false -->
```python
# Tests for Q7, worth 2.5 points
assert isinstance(abstract,bytes)
```

<!-- grade_id: cell-8673574466a3a00f, points: 2.5, grade: true, solution: false -->
```python
# Hidden tests for Q7, worth 2.5 points
```

## Q8
Finally, let's check if our abstract mentions 'blood pressure'.

1. First, turn the abstract into a string called `abstract_str`.
2. Then, create a boolean variable, `bp_bool` that reflects whether the abstract contains the string 'blood pressure'.

<!-- grade_id: cell-1b0903221728b832, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d5418c99702a206b, points: 2.5, grade: true, solution: false -->
```python
## Tests for Q8, 2.5 points
assert isinstance(bp_bool,bool)
assert isinstance(abstract_str,str)
```

<!-- grade_id: cell-b09fcf4aa42f980a, points: 2.5, grade: true, solution: false -->
```python
## Hidden tests for Q8, 2.5 points
```